# AVLTree

## Overview

In this notebook you will complete the following implementation of the balanced (AVL) binary search tree. Note that you should *not* be implementing the map-based API described in the plain (unbalanced) BSTree notebook — i.e., nodes in the AVLTree will only contain a single value.

<!-- grade_id: avltree, points: 0, grade: false, solution: true -->
```python
from IPython.core.debugger import Tracer

class AVLTree:
    class Node:
        def __init__(self, val, left=None, right=None):
            self.val = val
            self.left = left
            self.right = right

        def rotate_right(self):
           #print('rotating {} right around {}'.format(self.left.val, self.val))
            n = self.left
            self.val, n.val = n.val, self.val
            self.left, n.left, self.right, n.right = n.left, n.right, n, self.right
            
        def rotate_left(self):
           #print('rotating {} left around {}'.format(self.right.val, self.val))
            n = self.right
            self.val, n.val = n.val, self.val
            self.right, n.right, self.left, n.left = n.right, n.left, n, self.left
        
        @staticmethod
        def height(n):
            if not n:
                return 0
            else:
                return max(1+AVLTree.Node.height(n.left), 1+AVLTree.Node.height(n.right))

    def __init__(self):
        self.size = 0
        self.root = None
            
    @staticmethod
    def rebalance(t):
        if AVLTree.Node.height(t.left) > AVLTree.Node.height(t.right):
            if AVLTree.Node.height(t.left.left) >= AVLTree.Node.height(t.left.right):
                # left-left
               #print('left-left imbalance detected')
                t.rotate_right()
            else:
                # left-right
               #print('left-right imbalance detected')
                t.left.rotate_left()
                t.rotate_right()
        else:
            if AVLTree.Node.height(t.right.right) >= AVLTree.Node.height(t.right.left):
                # right-right
               #print('right-right imbalance detected')
                t.rotate_left()
            else:
                # right-left
               #print('right-left imbalance detected')
                t.right.rotate_right()
                t.rotate_left()
            
    def add(self, val):
        assert(val not in self)
        def add_rec(node):
            if not node:
                return AVLTree.Node(val)
            elif val < node.val:
                node.left = add_rec(node.left)
            else:
                node.right = add_rec(node.right)
            if abs(AVLTree.Node.height(node.left) - AVLTree.Node.height(node.right)) >= 2:
                AVLTree.rebalance(node)
            return node
        
        self.root = add_rec(self.root)
        self.size += 1
        
    def __delitem__(self, val):
        assert(val in self)
        def delitem_rec(node):
            if val < node.val:
                node.left = delitem_rec(node.left)
            elif val > node.val:
                node.right = delitem_rec(node.right)
            else:
                if not node.left and not node.right:
                    return None
                elif node.left and not node.right:
                    return node.left
                elif node.right and not node.left:
                    return node.right
                else:
                    # remove the largest value from the left subtree (t) as a replacement
                    # for the root value of this tree
                    t = node.left
                    if not t.right:
                        node.val = t.val
                        node.left = t.left                        
                    else:
                        n = t
                        while n.right.right:
                            n = n.right
                        t = n.right
                        n.right = t.left
                        node.val = t.val
            print('{} - {} == {}'.format(AVLTree.Node.height(node.left), AVLTree.Node.height(node.right),
                    abs(AVLTree.Node.height(node.left) - AVLTree.Node.height(node.right))))
            if abs(AVLTree.Node.height(node.left) - AVLTree.Node.height(node.right)) >= 2:
                AVLTree.rebalance(node)
            return node        
        
        self.root = delitem_rec(self.root)
        self.size -= 1
 
    def __contains__(self, val):
        def contains_rec(node):
            if not node:
                return False
            elif val < node.val:
                return contains_rec(node.left)
            elif val > node.val:
                return contains_rec(node.right)
            else:
                return True
        return contains_rec(self.root)
    
    def __len__(self):
        return self.size
    
    def __iter__(self):
        def iter_rec(node):
            if node:
                yield from iter_rec(node.left)
                yield node.val
                yield from iter_rec(node.right)
        yield from iter_rec(self.root)
        
    def pprint(self, width=128):
        """Attempts to pretty-print this tree's contents."""
        height = self.height()
        nodes  = [(self.root, 0)]
        prev_level = 0
        repr_str = ''
        while nodes:
            n,level = nodes.pop(0)
            if prev_level != level:
                prev_level = level
                repr_str += '\n'
            if not n:
                if level < height-1:
                    nodes.extend([(None, level+1), (None, level+1)])
                repr_str += '{val:^{width}}'.format(val='-', width=width//2**level)
            elif n:
                if n.left or level < height-1:
                    nodes.append((n.left, level+1))
                if n.right or level < height-1:
                    nodes.append((n.right, level+1))
                repr_str += '{val:^{width}}'.format(val=n.val, width=width//2**level)
        print(repr_str)
    
    def height(self):
        """Returns the height of the longest branch of the tree."""
        def height_rec(t):
            if not t:
                return 0
            else:
                return max(1+height_rec(t.left), 1+height_rec(t.right))
        return height_rec(self.root)
```

```python
t = AVLTree()
for x in [10, 5, 1]:
    print('adding: {}'.format(x))
    t.add(x)
    t.pprint()
```

Output:
```
adding: 10
                                                               10                                                               
adding: 5
                                                               10                                                               
                               5                                                               -                                
adding: 1
                                                               5                                                                
                               1                                                               10
```

<!-- grade_id: ll_test, points: 3, grade: true, solution: false -->
```python
# LL-fix (simple) test

from unittest import TestCase

def height(t):
    if not t:
        return 0
    else:
        return max(1+height(t.left), 1+height(t.right))

tc = TestCase()
t = AVLTree()

for x in [3, 2, 1]:
    t.add(x)
    
tc.assertEqual(height(t.root), 2)
tc.assertEqual([t.root.left.val, t.root.val, t.root.right.val], [1, 2, 3])
```

<!-- grade_id: rr_test, points: 3, grade: true, solution: false -->
```python
# RR-fix (simple) test

from unittest import TestCase

def height(t):
    if not t:
        return 0
    else:
        return max(1+height(t.left), 1+height(t.right))

tc = TestCase()
t = AVLTree()

for x in [1, 2, 3]:
    t.add(x)
    
tc.assertEqual(height(t.root), 2)
tc.assertEqual([t.root.left.val, t.root.val, t.root.right.val], [1, 2, 3])
```

<!-- grade_id: lr_test, points: 3, grade: true, solution: false -->
```python
# LR-fix (simple) test

from unittest import TestCase

def height(t):
    if not t:
        return 0
    else:
        return max(1+height(t.left), 1+height(t.right))

tc = TestCase()
t = AVLTree()

for x in [3, 1, 2]:
    t.add(x)
    
tc.assertEqual(height(t.root), 2)
tc.assertEqual([t.root.left.val, t.root.val, t.root.right.val], [1, 2, 3])
```

<!-- grade_id: rl_test, points: 3, grade: true, solution: false -->
```python
# RL-fix (simple) test

from unittest import TestCase

def height(t):
    if not t:
        return 0
    else:
        return max(1+height(t.left), 1+height(t.right))

tc = TestCase()
t = AVLTree()

for x in [1, 3, 2]:
    t.add(x)
    
tc.assertEqual(height(t.root), 2)
tc.assertEqual([t.root.left.val, t.root.val, t.root.right.val], [1, 2, 3])
```

<!-- grade_id: sort_test, points: 15, grade: true, solution: false -->
```python
# ensure key order is maintained after insertions and removals

from unittest import TestCase
import random

tc = TestCase()
vals = list(range(0, 100000000, 333333))
random.shuffle(vals)

t = AVLTree()
for x in vals:
    t.add(x)

for _ in range(len(vals) // 3):
    to_rem = vals.pop(random.randrange(len(vals)))
    del t[to_rem]

vals.sort()

for i,val in enumerate(t):
    tc.assertEqual(val, vals[i])
```

<!-- grade_id: balance_test, points: 15, grade: true, solution: false -->
```python
# stress testing

from unittest import TestCase
import random

tc = TestCase()

def traverse(t, fn):
    if t:
        fn(t)
        traverse(t.left, fn)
        traverse(t.right, fn)

def height(t):
    if not t:
        return 0
    else:
        return max(1+height(t.left), 1+height(t.right))
    
def check_balance(t):
    try:
        tc.assertLess(abs(height(t.left) - height(t.right)), 2, 'Tree is out of balance')
    except:
        Tracer()()

t = AVLTree()
vals = list(range(1000))
random.shuffle(vals)
for i in range(len(vals)):
    t.add(vals[i])
    print('added {}: {}'.format(i, vals[i]))
    for x in vals[:i+1]:
        tc.assertIn(x, t, 'Element added not in tree')
    traverse(t.root, check_balance)

random.shuffle(vals)
for i in range(len(vals)):
    del t[vals[i]]
    print('deleted {}: {}'.format(i, vals[i]))
    for x in vals[i+1:]:
        tc.assertIn(x, t, 'Incorrect element removed from tree')
    for x in vals[:i+1]:
        tc.assertNotIn(x, t, 'Element removed still in tree')
    traverse(t.root, check_balance)
```

Output:
```
added 0: 1
added 1: 799
added 2: 956
added 3: 669
added 4: 287
added 5: 777
added 6: 47
added 7: 180
added 8: 469
added 9: 64
added 10: 549
added 11: 478
added 12: 696
added 13: 569
added 14: 455
added 15: 873
added 16: 606
added 17: 782
added 18: 797
added 19: 654
added 20: 341
added 21: 754
added 22: 868
added 23: 438
added 24: 166
added 25: 344
added 26: 724
added 27: 218
added 28: 900
added 29: 615
added 30: 272
added 31: 554
added 32: 918
added 33: 412
added 34: 485
added 35: 626
added 36: 989
added 37: 237
added 38: 893
added 39: 73
added 40: 732
added 41: 519
added 42: 722
added 43: 590
added 44: 118
added 45: 370
added 46: 537
added 47: 858
added 48: 607
added 49: 636
added 50: 770
added 51: 660
added 52: 123
added 53: 29
added 54: 311
added 55: 450
added 56: 866
added 57: 993
added 58: 460
added 59: 833
added 60: 859
added 61: 928
added 62: 407
added 63: 788
added 64: 107
added 65: 579
added 66: 698
added 67: 314
added 68: 721
added 69: 294
added 70: 713
added 71: 773
added 72
```

```python

```