# Recursion

## Overview

For this assignment you will implement two recursive functions which both generate tree-shaped recursive processes, as a bit of a warmup for the upcoming binary tree data structure implementation.

## Tree recursive processes

### Ackermann's function

Ackermann's function is defined by the following recurrence:

$A(m, n) = 
\begin{cases}
n+1 & \mbox{if } m = 0 \\
A(m-1, 1) & \mbox{if } m > 0 \mbox{ and } n = 0 \\
A(m-1, A(m, n-1)) & \mbox{if } m > 0 \mbox{ and } n > 0.
\end{cases}$

You can read more about [Ackermann's function on Wikipedia](https://en.wikipedia.org/wiki/Ackermann_function).

Complete the following function, which should return the computed value of Ackermann's function for arguments `m` and `n`:

<!-- grade_id: ackermanss, points: 0, grade: false, solution: true -->
```python
def ackermanns(m, n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: ack_tests, points: 5, grade: true, solution: false -->
```python
from unittest import TestCase

tc = TestCase()

tc.assertEqual(ackermanns(0, 1), 2)
tc.assertEqual(ackermanns(1, 10), 12)
tc.assertEqual(ackermanns(2, 10), 23)
tc.assertEqual(ackermanns(3, 4), 125)
tc.assertEqual(ackermanns(4, 0), 13)
```

## Pascal's triangle

Pascal's triangle is a triangle arrangement of numbers where the top row contains the single number `1`, and the values in each following (centered) row are the sum of the value(s) in the row above. The following first five rows of Pascal's triangle should help demonstrate the idea:

              1
             1 1
            1 2 1
           1 3 3 1
          1 4 6 4 1
      
By convention, the rows and columns of Pascal's triangle are numbered starting from 0 — note that the 0th column of every row contains the value 1. To aid in the computation of edge cases (columns in rows that do not have two values above them), it is also convenient to imagine that there are columns in row 0 extending off in both directions that contain 0s. I.e., we might envision the first row of Pascal's triangle as follows:

    ... 0 0 0 1 0 0 0 ...
             1 1
            1 2 1
           1 3 3 1
          1 4 6 4 1

Wolfram Mathworld has a good writeup on the [properties and provenance of Pascal's Triangle](http://mathworld.wolfram.com/PascalsTriangle.html).

Complete the following function, which returns the value to be found in a given row and column of Pascal's triangle.

<!-- grade_id: pascals, points: 0, grade: false, solution: true -->
```python
def pascals(row, column):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# generate the first 10 rows of Pascal's Triangle
for row in range(10):
    print('{: ^45}'.format(' '.join(str(pascals(row, col)) for col in range(row+1))))
```

<!-- grade_id: pascal_tests, points: 5, grade: true, solution: false -->
```python
from unittest import TestCase

tc = TestCase()

tc.assertEqual(pascals(0, 0), 1)
tc.assertEqual(pascals(1, 0), 1)
tc.assertEqual(pascals(2, 1), 2)
tc.assertEqual(pascals(5, 1), 5)
tc.assertEqual(pascals(5, 2), 10)
tc.assertEqual(pascals(10, 5), 252)
```